![Noteable.ac.uk Banner](https://raw.githubusercontent.com/YusufNik/Noteable/refs/heads/main/images/Noteable%20NB%20Header%20Banner.png)

# Building Confidence with Noteable and Nbgrader: An Interactive Starter Notebook

## A friendly first notebook for running cells, fixing errors, and understanding grading

## Legend

<div style="border:1px solid #b8daff; background:#d9edf7; color:#0c5460; padding:16px; margin:8px 0; border-radius:4px;">
    In <b>blue</b>, the <b>instructions</b> and <b>goals</b> are highlighted. This tells you what we are trying to achieve.
</div>
<div style="border:1px solid #c3e6cb; background:#d4edda; color:#155724; padding:16px; margin:8px 0; border-radius:4px;">
    In <b>green</b>, key <b>information</b> and <b>concept explanations</b> are highlighted.
</div>
<div style="border:1px solid #ffeeba; background:#fff3cd; color:#856404; padding:16px; margin:8px 0; border-radius:4px;">
    In <b>yellow</b>, <b>exercises</b> and <b>tasks</b> are highlighted for you to try yourself.
</div>
<div style="border:1px solid #f5c6cb; background:#f8d7da; color:#721c24; padding:16px; margin:8px 0; border-radius:4px;">
    In <b>red</b>, <b>error interpretation</b> and <b>debugging tips</b> are highlighted.
</div>

## 1. Welcome to Noteable and Nbgrader

<div style="border:1px solid #b8daff; background:#d9edf7; color:#0c5460; padding:16px; margin:8px 0; border-radius:4px;">
    <b>Goal:</b> Learn the basic workflow used in Noteable notebooks and see how nbgrader manages different kinds of assessment cells.
</div>

### What you will do in this notebook
1. Run a few cells from top to bottom
2. Fix a small coding mistake
3. Load and inspect a toy dataset
4. Create a graph
5. Try an interactive widget
6. See how different nbgrader question types work

### How to work in the notebook
- Click inside a cell
- Run it with <code>Shift + Enter</code> or the Run button
- Read the output carefully
- If a cell is locked, you should not edit it
- If you see an error, use it as a clue, not a failure

```python
print("Starting setup: Importing libraries...")

# pandas helps us manage data
import pandas as pd

# matplotlib helps us draw graphs
import matplotlib.pyplot as plt

# ipywidgets allows us to create interactive sliders and buttons
import ipywidgets as widgets
from IPython.display import display, clear_output

print("Success! All libraries imported. You are ready to go!")
```

## 2. Understanding nbgrader question types

<div style="border:1px solid #c3e6cb; background:#d4edda; color:#155724; padding:16px; margin:8px 0; border-radius:4px;">
    <b>Core idea:</b> nbgrader lets instructors mix different types of assessment in one notebook.
</div>

### The main types you will see
- **Manually graded answer**  
  A human grader reads your response and assigns marks.

- **Manually graded task**  
  A task description cell that tells you what to do across one or more cells.

- **Autograded answer**  
  You write code, and hidden tests check whether it behaves correctly.

- **Autograded tests**  
  Read-only test cells that run during grading.

- **Hidden tests**  
  Extra tests hidden from students so they cannot see every grading detail.

- **AUTOTEST / HASHED AUTOTEST**  
  Short syntax that lets nbgrader generate tests automatically.

<div style="border:1px solid #ffeeba; background:#fff3cd; color:#856404; padding:16px; margin:8px 0; border-radius:4px;">
    <b>Tip:</b> The best student notebooks are short, visual, and predictable. Every section should tell the learner exactly what to do next.
</div>

<div style="border:1px solid #ffeeba; background:#fff3cd; color:#856404; padding:16px; margin:8px 0; border-radius:4px;">
    <b>Exercise:</b> Define the function <code>calculate_doubles</code> in the code cell below so it returns the value multiplied by 2.
</div>

<!-- grade_id: cell-c0728fb229a7c7d9, points: 0, grade: false, solution: true -->
```python
# Autograded Answer
# Write the function so that it returns double the input value.
def calculate_doubles(value):
    ### BEGIN SOLUTION
    return value * 2
    ### END SOLUTION

# Demonstration code for the function above.
# We set value to 10 so the first grading check is easy to understand.
value = 5
print(f'The double of {value} is {calculate_doubles(value)}')
```

<!-- grade_id: cell-5b9b91d1eb91b9c8, points: 1, grade: true, solution: false -->
```python
# [Grading Check] This locked cell automatically calculates your score upon submission.
### BEGIN HIDDEN TESTS
assert value == 10
assert calculate_doubles(value) == 20
### END HIDDEN TESTS
```

<div style="border:1px solid #ffeeba; background:#fff3cd; color:#856404; padding:16px; margin:8px 0; border-radius:4px;">
    <b>Exercise:</b> Change the code above so the value is <b>10</b> instead of 5. Press the <b><code>▶</code> Run</b> button again (or <code>Shift + Enter</code>) to see the new output!
</div>

## 3. Manually graded answer

<div style='border:1px solid #c3e6cb; background:#d4edda; color:#155724; padding:16px; margin:8px 0; border-radius:4px;'>
    <b>Manual answer:</b> This is the kind of response a human grader reads and marks.
</div>

**Question:** In one or two sentences, explain what the function <code>calculate_doubles</code> does and why returning a value is useful.

<div style='border:1px solid #ffeeba; background:#fff3cd; color:#856404; padding:16px; margin:8px 0; border-radius:4px;'>
    <b>Hint:</b> A good answer mentions the input, the output, and why the result can be reused in later calculations.
</div>

<b>Write your answer in the cell direct below </b> <i>(Click</i> `Enter` <i>to access and write-into the cell)</i> :

<!-- grade_id: cell-30c9b60dd2605181, points: 1, grade: true, solution: true -->
Student's answer here

## 4. Loading the Data

<div style="border:1px solid #b8daff; background:#d9edf7; color:#0c5460; padding:16px; margin:8px 0; border-radius:4px;">
    <b>Goal:</b> Load data so we have something to work with. Here, we will generate a standalone demo dataset.
</div>

```python
print("Fetching dataset... Please wait a moment.")

# Creating a simple standalone dataset for demonstration
data_dictionary = {
    'Category': ['A', 'B', 'C', 'D', 'E'],
    'Values': [10, 15, 7, 12, 18]
}
my_data = pd.DataFrame(data_dictionary)

print("Dataset loaded successfully! Here is a preview:")

# Display the top of the data
display(my_data.head())
```

## 5. Visualising the Results

<div style="border:1px solid #b8daff; background:#d9edf7; color:#0c5460; padding:16px; margin:8px 0; border-radius:4px;">
    <b>Goal:</b> Use our visual tools to map out what this data actually means.
</div>

```python
print("Generating visualization > Please wait a moment...")

# Set up the size of the graph
plt.figure(figsize=(6, 4))

# Create a simple bar chart
plt.bar(my_data['Category'], my_data['Values'], color='#4C72B0')
plt.title("Demo Data Visualization")
plt.xlabel("Category")
plt.ylabel("Values")

print("Visualization generated successfully! Check it out below:\n")

# Show the graph
plt.show()
```

## 6. Manually graded task

<div style='border:1px solid #ffeeba; background:#fff3cd; color:#856404; padding:16px; margin:8px 0; border-radius:4px;'>
    <b>Task:</b> Review the plot we've just created above and write a short note about what pattern you see in the values.
</div>

This is a task description cell. In a real assignment, the student would complete related work in one or more other cells, and a human grader would review the result.

<b>Write your answer in the cell direct below </b> <i>(Click</i> `Enter` <i>to access and write-into the cell)</i> :

Student's answer here

## 7. Error Interpretation & Debugging

<div style="border:1px solid #f5c6cb; background:#f8d7da; color:#721c24; padding:16px; margin:8px 0; border-radius:4px;">
    <b>Don't panic when you see red text!</b> Errors are just the computer asking for clarification. Learning to read them is a superpower. Let's look at a few common scenarios.
</div>

### A. Warnings (The Non-Disruptive Alerts)
Sometimes, Python gives you a heads-up that a tool is outdated. This generates a **Warning**, but your code will usually still run perfectly fine. *Note: We could hide these using `warnings.filterwarnings('ignore')`, but it's good practice to know they exist!*

```python
import warnings

# Simulating a deprecation warning
warnings.warn("This specific function is old and might be removed in future versions!", DeprecationWarning)

print("See? The code still ran successfully down here!")
```

### B. Import Errors (Missing Tools)
If you try to load a library that isn't installed on your environment, you will get a `ModuleNotFoundError`.

**How to fix it:** You can usually install it by typing `!pip install library_name` in a code cell, or `pip install library_name` in the terminal from the launcher. 

**Noteable Users:** If a core library you need is missing from this environment, please contact the EDINA Noteable team! They will look into updating the Notebook server presets for everyone.

```python
print("Attempting to import a library that doesn't exist...")

# This will trigger an error!
import super_secret_fake_library_for_science
```

### C. Isolated Syntax Errors
Sometimes you just make a typo. Below, we forgot a closing parenthesis. This cell will fail, but notice how it doesn't break the entire notebook. 

<div style="border:1px solid #ffeeba; background:#fff3cd; color:#856404; padding:16px; margin:8px 0; border-radius:4px;">
    <b>Fix It:</b> Add the missing <code>)</code> at the end of the print statement, then run the cell again!
</div>

```python
# Run this cell to see the SyntaxError, then fix it!
print("I am missing a closing bracket"
```

<!-- grade_id: cell-1394d22dd6af1431, points: 1, grade: true, solution: false -->
```python
# [Grading Check] This locked cell automatically calculates your score upon submission.
### BEGIN HIDDEN TESTS
# Dynamically look at the code executed immediately before this test cell
student_code = In[-2] if len(In) >= 2 else ""
# This type of testing is supposedly unstable/fragile - use with caution

# Define the exact core phrase that MUST execute to prove they fixed it
target_phrase = 'print("I am missing a closing bracket")'

# If the text is missing, or they didn't fix the syntax, throw an error
if target_phrase not in student_code:
    raise AssertionError("The syntax error has not been fixed or the cell was not executed.")
### END HIDDEN TESTS
```

### D. Cascading Errors
Sometimes an error in one cell prevents important data from being created, which causes the *next* cell to fail. Below, we try to divide by zero, which stops the variable `important_calculation` from ever being saved.

<div style='border:1px solid #b8daff; background:#d9edf7; color:#0c5460; padding:16px; margin:8px 0; border-radius:4px;'>
    <b>Debugging rule:</b> Fix the first error first. Many later errors disappear automatically.
</div>

```python
# This cell fails because we can't divide by zero.
important_calculation = 100 / 0

print("This line will never print, because the error stopped the code above it.")
```

Because the cell above failed, the cell below will give a `NameError`—it has no idea what `important_calculation` is!

<div style="border:1px solid #ffeeba; background:#fff3cd; color:#856404; padding:16px; margin:8px 0; border-radius:4px;">
    <b>Fix It:</b> Go back to the cell above and change <code>100 / 0</code> to <code>100 / 2</code>. Run it, and then run the cell below again.
</div>

```python
# This fails until the cell above is fixed!
final_result = important_calculation + 50
print(f"The final result is: {final_result}")
```

<!-- grade_id: cell-4da7e00c1ea55667, points: 1, grade: true, solution: false -->
```python
# [Grading Check] This locked cell automatically calculates your score upon submission.
### BEGIN HIDDEN TESTS
isinstance(final_result, (int, float))
### END HIDDEN TESTS
```

## 8. Interactive Explorations

<div style="border:1px solid #c3e6cb; background:#d4edda; color:#155724; padding:16px; margin:8px 0; border-radius:4px;">
    <b>Interactivity:</b> Many libraries (like mapping or 3D tools) have interactivity built right in! For basic data, we can also use sliders to manipulate our graphs dynamically.
</div>

Run the cell below, and try using the controls to see how the output updates in real-time.

```python
print("Building interactive widget...")

def interactive_function(multiplier):
    # Scale the data based on the slider
    updated_values = my_data['Values'] * multiplier
    
    plt.figure(figsize=(6, 4))
    plt.bar(my_data['Category'], updated_values, color='#55A868')
    plt.title(f"Values Multiplied by {multiplier}")
    plt.xlabel("Category")
    plt.ylabel("Scaled Values")
    plt.ylim(0, 100) # Lock the Y-axis so it doesn't jump around
    plt.show()

# Create the interactive slider
my_widget = widgets.IntSlider(
    value=1,
    min=1,
    max=5,
    step=1,
    description='Multiplier:',
    continuous_update=False
)

print("Widget ready! Use the controls below:\n")

# Display the widget linked to our function
widgets.interact(interactive_function, multiplier=my_widget);
```

## Take-away

<div style='border:1px solid #b8daff; background:#d9edf7; color:#0c5460; padding:16px; margin:8px 0; border-radius:4px;'>
    Congratulations! You've successfully completed this module. You've learned how to handle errors natively, read tracebacks, and interact with live data. Keep experimenting and having fun with Python, Noteable, and nbgrader!
</div>

### You now know how to:
- run a Noteable notebook from top to bottom
- recognise different cell types
- understand autograded and manually graded work
- debug common notebook errors
- use simple data, visualisation, and widgets

![Noteable license](https://raw.githubusercontent.com/YusufNik/Noteable/refs/heads/main/images/Noteable%20Notebook%20Footer.png)